# Practice Assignment 2

In this assignment, you will explore the Online Shoppers Purchasing Intention Dataset to build a machine learning model that predicts whether a visitor will make a purchase. you will analyze data in the online_shoppers_intention.csv file.You'll use functions from the pandas module for loading, inspecting and querying the data. You are expected to preprocess the data and use the data to train machine learning models from the sklearn library to solve problems.

The useful dataset link: https://archive.ics.uci.edu/dataset/468/online+shoppers+purchasing+intention+dataset

For each question, there are clear instructions in each cell. Follow those instructions and write the code after each block of:

YOUR CODE HERE

Please use the exact variable name if it is specified in the comment.

We’ll run a Python test script against your program to test whether each function implementation is correct.

<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
%%capture
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################
"""
Import all libraries needed for the entire exercise: numpy, pyplot, and seaborn
"""
import imp
import numpy as np
from nose.tools import assert_equal
from pandas.util.testing import assert_frame_equal, assert_series_equal
sol = imp.load_compiled("sol", "./.sol.py")
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
"""
Part 01: Setup and Data Loading

1. Import the pandas library.

2. Use pandas to read the dataset from the .csv file and save it to "shoppers_df". 

3. In certain datasets, a value of 0 (zero) can indicate missing information that was not recorded. 
   In the case of our dataset, assume that 0's represent missing data for the following analysis.
   
   Replace all 0 or 0.0 values in the DataFrame with NaN using replace(). 
   Hint: Make sure to set inplace to True.

3. Display the first 5 rows of the DataFrame "shoppers_df".

"""

# YOUR CODE HERE

```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_frame_equal(shoppers_df, sol.shoppers_df)

```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
"""
Part 02: Data Exploration

1. Display the shape of the DataFrame "shoppers_df" to know the number of rows and columns. 
   Store the shape in a variable called "shape_of_df"

2. Use the `info()` method on "shoppers_df" to display the summary including the data types 
   and the presence of null values.

3. Create a variable named "missing_values_count" that stores the total count of missing 
   values in the DataFrame.
"""

# YOUR CODE HERE

```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(shape_of_df, sol.shape_of_df)

```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
"""
Part 03: Data Pre-processing

1. First save old df to "new_shoppers_df"

2. Inspect the dataset for missing values in each column and handle them appropriately. 
   For each column, fill any missing values with the mode of that column. Implement this using a for loop.

3. Drop the 'Month' column from "new_shoppers_df" as it will not be used in further analysis.

4. Perform one-hot encoding on the 'VisitorType' column. 
  
  First, use `pd.get_dummies()` to create a DataFrame "visitor_dummies". 
  
  Then, concatenate "visitor_dummies" with "shoppers_df"
  
  a. Hint: For more about pd.get_dummies(): 
       https://www.geeksforgeeks.org/python-pandas-get_dummies-method/
  b. Do not forget to drop the old 'VisitorType' column
  c. Set parameter drop_first = True. We only need k-1 catergoies to represent
     k catergorical values. 

5. After the encoding and concatenation, display the first 5 rows of "new_shoppers_df" to verify the changes.
  
"""

# YOUR CODE HERE

```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_frame_equal(new_shoppers_df, sol.new_shoppers_df)

```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
"""
Part 04: Data Preparation for Modeling

1. Separate the dataset into features and target variable ("Revenue"), 
   storing them in "X" and "y" respectively.
   
2. Further, split the dataset into training and test sets using `train_test_split()` 
   with a test size of 0.2 and a random state of 42. 
   Store the outputs in "X_train", "X_test", "y_train", "y_test".

"""

# YOUR CODE HERE

```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_frame_equal(X, sol.X)
assert_series_equal(y, sol.y)

assert_frame_equal(X_train, sol.X_train)
assert_frame_equal(X_test, sol.X_test)
assert_series_equal(y_train, sol.y_train)
assert_series_equal(y_test, sol.y_test)

```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
"""
Part 05: Model Selection & Training

1. First, import libraries: LogisticRegression, accuracy_score and StandardScaler. 
   Use StandardScaler() to Scale the features
   
2. Train a logistic regression model using the training data and name the model "logistic_model". 
   Set max_iter to 1000.

3. Make predictions on the test set using the "logistic_model" 
   and store the predictions in a variable "y_pred".

4. Calculate the accuracy of the model on the test set and store it 
   in a variable "model_accuracy".

5. Display the result for "model_accuracy"
"""

# YOUR CODE HERE

```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

np.testing.assert_equal(y_pred, sol.y_pred)
assert_equal(model_accuracy, sol.model_accuracy)

```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
"""
Part 06: Visualize the Confusion Matrix

1. Import matplotlib.pyplot and seaborn here

2. Calculate the confusion matrix using the actual and the predicted values.

3. Create a heatmap using `matplotlib` and `seaborn` to visualize the confusion matrix.

4. Ensure to label the heatmap with actual labels for better readability.

NOTE: 
- Your plot cannot be autograded.  
- If you go to the guide, there will be a PDF with the solution plot.
- Please check it and if your answer matches, set the variable q6_plot_check = 'yes' .   

"""

q6_plot_check = 'no' # change to yes after you verify your plot

# YOUR CODE HERE

```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal('yes', q6_plot_check)
```

```python

```