# Practice Assignment 1

In this assignment, you will analyze data in the Air Quality dataset to build a machine learning model that predicts air quality. You will analyze data in the ```AirQualityUCI.xlsx``` file and use functions from the pandas module for loading, inspecting, and querying the data. You are expected to preprocess the data and use the data to train machine learning models from the sklearn library to solve problems.

More information about the dataset: https://archive.ics.uci.edu/dataset/360/air+quality

For each question, there are clear instructions in each cell. Follow those instructions and write the code after each block of:
```
# YOUR CODE HERE
```

Please use the exact variable name if it is specified in the comment.

We’ll run a Python test script against your program to test whether each function implementation is correct.

<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
%%capture
########################################################
### EXECUTE THIS CELL BEFORE YOU TEST YOUR SOLUTIONS ###
########################################################
"""
Import all libraries needed for the entire exercise: numpy, pyplot
"""
import imp
import numpy as np
from nose.tools import assert_equal
from pandas.util.testing import assert_frame_equal, assert_series_equal
sol = imp.load_compiled("sol", "./.sol.py")
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
"""
Part 01: Load the Dataset

1. Import the pandas library.

2. Use pandas to read the dataset from xlsx file. and save them to "air_quality_df". 
   read_excel() will be a good tool to read it (Hint: You will need to set the "engine" parameter to engine='openpyxl').

3. Display the first 5 rows of the DataFrame.

"""

# YOUR CODE HERE

```

Output:
```
'\nPart 01: Load the Dataset\n\n1. Import the pandas library.\n\n2. Use pandas to read the dataset from xlsx file. and save them to "air_quality_df". \n   read_excel() will be a good tool to read it (Hint: You will need to set the "engine" parameter to engine=\'openpyxl\').\n\n3. Display the first 5 rows of the DataFrame.\n\n'
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

try:
    assert_frame_equal(air_quality_df, sol.air_quality_df)
    print("Pass")
except AssertionError as e:
    print("Test failed:", e)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-3-5cdd61478037> in <module>()
      4 
      5 try:
----> 6     assert_frame_equal(air_quality_df, sol.air_quality_df)
      7     print("Pass")
      8 except AssertionError as e:

NameError: name 'air_quality_df' is not defined
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
"""
Part 02: Data Cleaning

Data cleaning is crucial to ensure the quality of a linear regression model's prediction(s)
1. Create a new dataframe called "new_air_quality_df" and set it equal to "air_quality_df"

2. Drop the rows having missing values in "new_air_quality_df"

3. Identify and remove any non-numeric or irrelevant columns that may interfere with model training. Let's 
assume 'Date' and 'Time' are irrelevant for our regression task.
"""

# YOUR CODE HERE

```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

try:
    assert_frame_equal(new_air_quality_df, sol.new_air_quality_df)
    print("Pass")
except AssertionError as e:
    print("Test failed:", e)

```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
"""
Part 03: Exploratory Data Analysis (EDA)

Perform basic EDA to understand the data in "new_air_quality_df"
1. Import the matplotlib and seaborn libraries

2. Plot histograms for the various features using matplotlib
   a. Set the bins to 15
   b. Set the figure size to (15, 10)
   
3. Plot a heatmap of the correlation matrix using seaborn
   a. Set the annotations to True
   b. Set the figure size to (12, 10)
   
   Note he: 
      Seaborn is a simple data visualization library based on matplotlib
      Here is a useful link for plotting Heatmap using seaborn: https://www.geeksforgeeks.org/how-to-create-a-seaborn-correlation-heatmap-in-python/
   
   
NOTE: 
- Your plot cannot be autograded.  
- If you go to the guide, there will be a PDF with the solution plot.
- Please check it and if your answer matches, set the variable q3_plot_check = 'yes' .   

"""
q3_plot_check = 'no' # change to yes after you verify your plot

# YOUR CODE HERE

```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

try:
    assert_equal('yes', q3_plot_check)
    print("Pass")
except AssertionError as e:
    print("Test failed:", e)


```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
"""
Part 04: Feature Engineering and Selection
Prepare your data for the linear regression model
1. Select the target variable, which should be the 'NO2(GT)' column
2. Select the features, which will be all columns except for the target variable
3. Split your data into training and testing sets using an 80-20 split and a random state of 42. The variables
should be "X_train", "X_test", "y_train", and "y_test".
4. Standardize your features using StandardScaler()

"""
# YOUR CODE HERE

```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

try:
    assert_frame_equal(X, sol.X)
    assert_series_equal(y, sol.y)

    np.testing.assert_equal(X_train, sol.X_train)
    np.testing.assert_equal(X_test, sol.X_test)
    assert_series_equal(y_train, sol.y_train)
    assert_series_equal(y_test, sol.y_test)
    print("Pass")
except AssertionError as e:
    print("Test failed:", e)


```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
"""
Part 05: Build the Linear Regression Model

Now, it's time to build and train your linear regression model
1. Create a LinearRegression object

2. Fit your model to the training data

3. Predict the target variable for the test set. Save the result to "y_pred" and display it.
"""

# YOUR CODE HERE

```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

try:
    np.testing.assert_equal(y_pred, sol.y_pred)
    print("Pass")
except AssertionError as e:
    print("Test failed:", e)



```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
"""
Part 06: Model Evaluation

Evaluate the performance of your linear regression model using the test set

  1. Calculate the Mean Squared Error (MSE). Save it into "mse" variable.
"""
# YOUR CODE HERE

```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

try:
    np.testing.assert_almost_equal(mse, sol.mse, 2)
    print("Pass")
except AssertionError as e:
    print("Test failed:", e)



```

```python

```