# Practice Assignment 2
# Joining & Filtering Data

For this practice assignment some of the code has been written for you ahead of time.  Cells you just need to run are mark 'Run this cell'.  For each question, there are clear instructions in each cell. Follow those instructions and write the code after each block of:

YOUR CODE HERE

Please use the exact variable name if it is specified in the comment.

We’ll run a Python test script against your program to test whether each function implementation is correct.

<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
%%capture
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################
"""
Loading the necessary python libraries
"""
import imp
import pandas as pd
import numpy as np
from nose.tools import assert_equal
from pandas.util.testing import assert_frame_equal, assert_series_equal
sol = imp.load_compiled("sol", "./.sol.py")
```

<!-- grade_id: init_1, points: 0, grade: true, solution: false -->
```python
"""
First we need to load the dataset and preprocess the data. The below code comes from some of the work done in Practice Assignment 01.
"""

#Run this cell
restaurant_location_df = pd.read_csv("geoplaces2.csv")
restaurant_accept_df = pd.read_csv("chefmozaccepts.csv")
restaurant_cuisine_df = pd.read_csv("chefmozcuisine.csv")
rating_final_df = pd.read_csv("rating_final.csv")
user_profile_df = pd.read_csv("userprofile.csv")

restaurant_location_df = restaurant_location_df[["placeID", "name", "address", "city", "state", "country", "price", "alcohol"]]
restaurant_location_df = restaurant_location_df.apply(lambda x: x.astype(str).str.lower())
restaurant_location_df = restaurant_location_df.astype({'placeID':'int',"name":"string", "address":"string","city":"string", "state":"string", "country":"string", "price":"string", "alcohol": "string"})
restaurant_location_df = restaurant_location_df.replace('\?', np.nan, regex=True).dropna().reset_index(drop = True)
user_profile_df = user_profile_df.replace('\?', np.nan, regex=True)
user_profile_df = user_profile_df[['userID', 'smoker', 'drink_level', 'marital_status', 'birth_year']]

#Since The values in column 'smoker' are strings, we have to convert them to boolean values
user_profile_df['smoker'] = user_profile_df['smoker'].map({'true': True, 'false': False})
###
```

<!-- grade_id: init_2, points: 0, grade: true, solution: false -->
```python
"""
Lets first examine the columns and datatypes for our "restaurant_location_df" dataframe.
"""

#Run this cell
restaurant_location_df.info()
```

<!-- grade_id: init_3, points: 0, grade: true, solution: false -->
```python
"""
Now let's display the "restaurant_location_df" dataframe.  As a reminder, we can use display() to display
dataframes as easy-to-read html tables.
"""

#Run this cell
display(restaurant_location_df)
```

<!-- grade_id: init_4, points: 0, grade: true, solution: false -->
```python
"""
To quickly get an idea of the rows in all of the dataframes, we'll use .head() in the following 4 cells.
View the first 5 rows for "restaurant_accept_df".
"""

#Run this cell
restaurant_accept_df.head()
```

<!-- grade_id: init_5, points: 0, grade: true, solution: false -->
```python
"""
View the first 5 rows for "restaurant_cuisine_df".
"""

#Run this cell
restaurant_cuisine_df.head()
```

<!-- grade_id: init_6, points: 0, grade: true, solution: false -->
```python
"""
View the first 5 rows for "user_profile_df".
"""

#Run this cell
user_profile_df.head()
```

<!-- grade_id: init_7, points: 0, grade: true, solution: false -->
```python
"""
View the first 5 rows for "rating_final_df".
"""

#Run this cell
rating_final_df.head()
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
"""
Part 01

1. We want to see which restaurants can accept VISA payments.  In order to do this we need to merge
"restaurant_location_df" and "restaurant_accept_df".  Do this as a left join and store in a new dataframe called
"joined_df".  Examine both dataframes to identify the best columns to merge on.

2. Display the first 5 rows of "joined_df"

3. Create a new dataframe called "restaurant_visa_df" which contains only two columns from "joined_df":
'name' and 'Rpayment'.  Also filter the dataframe so that the 'Rpayment' column only contains those who paid
with 'VISA'. 

*** Please note: After performing the merge and filtering operations, ensure reset_index(drop=True) 
before submission to aviod Index-related Errors.

4. Display the first 10 rows of "restaurant_visa_df"
"""

# YOUR CODE HERE

```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_frame_equal(joined_df, sol.joined_df)
print("test joined_df correct")
assert_frame_equal(restaurant_visa_df, sol.restaurant_visa_df)
print("test restaurant_visa_df correct")

```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
"""
Part 02

1. First lets merge "restaurant_location_df" and "rating_final_df" with a left join.  Store the result in a new
dataframe called "restaurant_mean_rating_df".  Examine both dataframes to identify the best columns to merge on.

2. "restaurant_mean_rating_df" should only keep 3 columns after it has been merged: 'name', 'food_rating',
and 'service_rating'.

3. Filter "restaurant_mean_rating_df" by its 'name' column, for restaurants with the name
'restaurant bar coty y pablo'

4. Display the updated "restaurant_mean_rating_df".  Hint: If everything was completed correctly, you should
have 3 columns and 6 rows.

5. Calculate the average (mean) of the 'food_rating' column and store in a new variable called
"mean_food_rating_pablo"

6. Calculate the average (mean) of the 'service_rating' column and store in a new variable called
"mean_ser_rating_pablo"

7. Print both of the means you just calculated
"""

# YOUR CODE HERE

```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################


assert_frame_equal(restaurant_mean_rating_df, sol.restaurant_mean_rating_df)
print("test restaurant_mean_rating_df correct")
assert_equal(mean_food_rating_pablo, sol.mean_food_rating_pablo)
print("test mean_food_rating_pablo correct")
assert_equal(mean_ser_rating_pablo, sol.mean_ser_rating_pablo)
print("test mean_ser_rating_pablo correct")


```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
"""
Part 03

For this part we want to see the average rating of all restaurants by "single smokers" and the average
rating of all restaurants by "married non-smokers".

1. Merge "rating_final_df" (as the left df) and "user_profile_df" (as the right df) using an inner join, and name the resulting dataframe
"join_rating_df".  Examine both dataframes to identify the best columns to merge on.

2. Create a new dataframe called "single_smoke_df", which has 3 columns: 'smoker', 'marital_status',
and 'rating'. This should be a slice of the join_rating_df your just created.

3. Filter "single_smoke_df" where the 'smoker' column is True, and the 'marital_status' column is 'single'

4. Calcualate the average (mean) of the ratings in "single_smoke_df" and store in a variable called
"single_smoke_rating"

5. Repeat steps 2 - 4 for married non-smokers and name the dataframe "married_nonsmoke_df" and the final average
variable "married_nonsmoke_rating".  Hint: We are asking for non-smokers now, so adjust your filters as needed.

6. Print out the results for your variables rounded to 2 decimal places ('single_smoke_rating' and
'married_nonsmoke_rating')
"""

# YOUR CODE HERE

```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_frame_equal(join_rating_df, sol.join_rating_df)
print("test join_rating_df correct")

assert_frame_equal(single_smoke_df, sol.single_smoke_df)
print("test single_smoke_df correct")

assert_equal(single_smoke_rating, sol.single_smoke_rating)
print("test single_smoke_rating correct")

assert_frame_equal(married_nonsmoke_df, sol.married_nonsmoke_df)
print("test married_nonsmoke_df correct")
assert_equal(married_nonsmoke_rating, sol.married_nonsmoke_rating)
print("test married_nonsmoke_rating correct")

```